evaluation_interval: 1m

rule_files:
  - api_server_availability_recording_rules.yaml

tests:
  # Start of a new test case. The interval and input series must be under this item.
  - name: Test KonfluxAPIServerHighErrorRate
    interval: 1m # Correct placement and spelling of interval
    input_series:
      - series: 'code:apiserver_request_total:rate5m{code="503"}'
        values: '0+5x60'
      - series: 'code:apiserver_request_total:rate5m{code="200"}'
        values: '0+95x60'

    promql_expr_test:
      - expr: konflux_up
        eval_time: 1m
        exp_samples:
          - labels: '{__name__="konflux_up", service="api-server"}'
            value: 1

  - name: Test api_server error_rate_10m at 5 percent (5xx errors)
    interval: 1m
    input_series:
      - series: 'apiserver_request_total{code="503", source_cluster="test-cluster"}'
        values: '0+5x600'
      - series: 'apiserver_request_total{code="200", source_cluster="test-cluster"}'
        values: '0+95x600'

    promql_expr_test:
      - expr: api_server:error_rate_10m
        eval_time: 10m
        exp_samples:
          - labels: '{__name__="api_server:error_rate_10m", service="api-server", source_cluster="test-cluster"}'
            value: 5

  - name: Test api_server error_rate_10m at 10 percent (4xx + 5xx errors)
    interval: 1m
    input_series:
      - series: 'apiserver_request_total{code="503", source_cluster="test-cluster-mixed"}'
        values: '0+5x600'
      - series: 'apiserver_request_total{code="404", source_cluster="test-cluster-mixed"}'
        values: '0+5x600'
      - series: 'apiserver_request_total{code="200", source_cluster="test-cluster-mixed"}'
        values: '0+90x600'

    promql_expr_test:
      - expr: api_server:error_rate_10m
        eval_time: 10m
        exp_samples:
          - labels: '{__name__="api_server:error_rate_10m", service="api-server", source_cluster="test-cluster-mixed"}'
            value: 10

  - name: Test api_server error_rate_30d_avg with constant rate
    interval: 1h
    input_series:
      - series: 'api_server:error_rate_10m{service="api-server", source_cluster="test-cluster-avg"}'
        values: '10+0x720'

    promql_expr_test:
      - expr: api_server:error_rate_30d_avg
        eval_time: 720h
        exp_samples:
          - labels: '{__name__="api_server:error_rate_30d_avg", service="api-server", source_cluster="test-cluster-avg"}'
            value: 10

  - name: Test api_server error_rate_30d_stddev with constant rate
    interval: 1h
    input_series:
      - series: 'api_server:error_rate_10m{service="api-server", source_cluster="test-cluster-stddev"}'
        values: '10+0x720'

    promql_expr_test:
      - expr: api_server:error_rate_30d_stddev
        eval_time: 720h
        exp_samples:
          - labels: '{__name__="api_server:error_rate_30d_stddev", service="api-server", source_cluster="test-cluster-stddev"}'
            value: 0
