rule_files:
  - prometheus.caching_squid_alerts.yaml

tests:
  # Test CachingSquidDown alert - FIRING
  - name: CachingSquidDownFiring
    interval: 1m
    input_series:
      - series: 'konflux_up{namespace="proxy", service="squid"}'
        values: '0x10'

    alert_rule_test:
      - eval_time: 7m
        alertname: CachingSquidDown
        exp_alerts:
          - exp_labels:
              severity: critical
              slo: "true"
              namespace: "proxy"
              service: "squid"
            exp_annotations:
              summary: "Caching Squid proxy is down"
              description: "Squid exporter cannot scrape the Squid proxy successfully in namespace proxy. The proxy has been down for 5 minutes."
              alert_team_handle: "<!subteam^S07NDQV6A4D>"
              team: vanguard
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/vanguard/sre/caching-squid-down.md

  # Test CachingSquidDown alert - NOT FIRING
  - name: CachingSquidDownNotFiring
    interval: 1m
    input_series:
      - series: 'konflux_up{namespace="proxy", service="squid"}'
        values: '1x10'

    alert_rule_test:
      - eval_time: 7m
        alertname: CachingSquidDown
        exp_alerts: []
